# Week 2 Exercises

This week, in modules 9 through 14 we worked on understanding how to control the flow of code execution through things like loops (for, while, range) and conditions (if).  See what you learned.

### 9.1 - An IF statement in Python can:

> A) Only perform one conditional test
>
> B) Returns the indented actions if the test is True, or None if it is not True
>
> C) Can contain as many branches of logic as the programmer wants
>
> D) Executes the next line of code only if the test is True, then moves to always execute the remaining code

```python
answer91 = "C"
```

### 9.2 - What will the following code print to the screen?

```python
age = 21

if age > 18:
    label = 'adult'
label = 'kid'

print(label)
```

> A) `kid`
>
> B) `adult`
>
> C) `adult` (new line) `kid`
>
> D) _nothing_

```python
answer92 = "A"
```

### 9.3 - What will the follow code print to the screen?

```python
temp = 22
precipitation = 'wet'

if precipitation != 'wet':
    weather = 'dry'
else:
    if temp > 32:
        weather = 'rain'
    elif temp < 32:
        weather = 'sleet'
    elif temp < 28:
        weather = 'mix'
    else:
        weather = 'snow'

print(weather)
```

> A) `sleet`
>
> B) `mix`
>
> C) `now`
>
> D) `dry'

```python
answer93 = "A"
```

### 10.1 - How many times will the following code print the phrase 'Hello World!'?

```python
start = 1
stop = 1000

for num in range(start, stop):
    print('Hello World!')
    start = start + 5
```

> A) 200 times
>
> B) 199 times
>
> C) 999 times
>
> D) 201 times

```python
answer101 = "C"
```

### 10.2 - Which of these best describes how the `range` function works?

> A) It generates a list of values that the for loop can then iterate over.
>
> B) It pre-generates the code required to execute the for loop according to the range() parameters
>
> C) If just runs the specified code as many times as necessary to create the right output.
>
> D) It returns an iterable object that returns the next value for each time through the loop.


```python
answer102 = "A"
```

### 11.1 - If you accidentally run an infine while loop in Jupyter:

> A) Close the browser window and it will automatically stop.
>
> B) Interrupt or restart the kernel.
>
> C) You will overheat the server and burn it up.
>
> D) Slack the professor so that he can stop the loop for you.

```python
answer111 = "B"
```

### 11.2 - What will the following code print to the screen?

```python
counter = 1

while (counter < 10):
    print(counter)
    break
```

> A) `1`
>
> B) `1,2,3,4,5,6,7,8,9`
>
> C) `1,2,3,4,5,6,7,8,9,10`
>
> D) _nothing_

```python
answer112 = "A"
```

### 11.3 - What will the follow code do?

```python
status = True
while (status):
    status =  False
while (status):
    status = True
    print("The status is now true")
```

> A) Only the first line of code will execute
>
> B) The code will print 'The status is now true'
>
> C) The code will all run but nothing will print
>
> D) None of these

```python
answer113 = "C"
```

### 12.1 - Which of the following is **not** a correct statement about the `range` function?

> A) `range()` always produces incrementing numbers
>
> B) `range()` doesn't actually create the whole list of output numbers all at once
>
> C) `range()` can only produce integer values
>
> D) `range()` can use negative numbers for start and end values

```python
answer121 = "A"
```

### 12.2 - How many numbers will the following code output?

```python
for value in range(-1,1,1):
    print(value)
```

> A) _None_
>
> B) 1
>
> C) 2
>
> D) 3

```python
answer122 = "C"
```

### 12.3 - What is the last number this `range` call will produce?

```python
for value in range(8,10,-2):
    print(value)
```

> A) 8
>
> B) 10
>
> C) 9
>
> D) _nothing_

```python
answer123 = "D"
```

### 13.1 - What is another name for the inline if statement?

> A) Ternary operator
>
> B) Conditional
>
> C) Three-part if
>
> D) None of these

```python
answer131 = "A"
```

### 13.2 - What does this code output?

```python
age = 67
gender = 'male'

print(
    ('young ' if age < 18 else ('retired ' if age > 65 else '')) + 
    ('lady' if gender == 'female' else 'guy')
)
```

> A) `man`
>
> B) `retired guy`
>
> C) `retired`
>
> D) _nothing_

```python
answer132 = "B"
```

### 14.1 NAND

Inside of computer chips, it's often easier to do a logical operation called NAND (not and).  Write a function that takes two boolean parameters and returns the NAND of those values.  Here is the k-map for NAND:

```
NAND: 

        TRUE   FALSE
TRUE    False  True
FALSE   True   True
```


<!-- grade_id: average-answer, points: 0, grade: false, solution: true -->
```python
def nand(a, b):
    if a == 1 and b == 1:
        return False
    else:
        return True
        
    ### BEGIN SOLUTION
    answer = None
    ### END SOLUTION
    
    return answer
```

```python
nand(1,1)
```

Output:
```
False
```

<!-- grade_id: average-tests, points: 5, grade: true, solution: false -->
```python
assert nand(True, True) == False
assert nand(True, False) == True
assert nand(False, True) == True
assert nand(False, False) == True
```

### 14.2 LACE Score
There is a simple readmission index called the LACE Score: https://www.hindawi.com/journals/bmri/2015/169870/tab1/

Use this documentation to create a function that can compute LACE scores based on the 4 input parameters.  Below is the function signature and documentation that you start with.

```python
def LOS_score(length_of_stay):
        if length_of_stay < 1:
            return 0
        elif length_of_stay == 1:
            return 1
        elif length_of_stay == 2:
            return 2
        elif length_of_stay == 3:
            return 3
        elif length_of_stay >= 4:
            return 4
        elif length_of_stay >= 7:
            return 5
        else:
            return 7
        
def acute_flag_score(acute_flag):
        if acute_flag:
            return 3
        else:
            return 0 
    
def charlson_score(charlson):
        if charlson == 0:
            return 0
        elif charlson == 1:
            return 1
        elif charlson == 2:
            return 2
        elif charlson == 3:
            return 3
        else:
            return 5
    
def ed_visits_score(ed_visits):
        if ed_visits == 0:
            return 0
        elif ed_visits == 1:
            return 1
        elif ed_visits == 2:
            return 2
        elif ed_visits == 3:
            return 3
        else:
            return 4
        
```

<!-- grade_id: bmi-answer, points: 0, grade: false, solution: true -->
```python
def LACE(length_of_stay, acute_flag, charlson, ed_visits):
    
    """(int, bool, int, int) -> int
    This function uses the logic from https://www.hindawi.com/journals/bmri/2015/169870/tab1/
    to compute the LACE score for this patient.
    
    >>> LACE(4, False, 1, 0)
    5
    
    >>> LACE(4, True, 4, 7)
    16
    
    """
    ### BEGIN SOLUTION
    
    score = LOS_score(length_of_stay) + \
        acute_flag_score(acute_flag) + \
        charlson_score(charlson) + \
        ed_visits_score(ed_visits)
    
    ### END SOLUTION
    
    return score
```

```python
LACE(4, False, 1, 0)
```

Output:
```
5
```

```python
LACE(4, True, 4, 7)
```

Output:
```
16
```

<!-- grade_id: bmi-test, points: 10, grade: true, solution: false -->
```python
assert LACE(4, False, 1, 0) == 5
assert LACE(4, True, 4, 7) == 16
```

### 14.3 Care Management Criteria

Care managers use LACE as part of the criteria for assigning a care coodinator to a patient who has been recently discharged. If the score is above 10, then a care coordinator will be assigned. The other criteria they use is if the patient has been discharged with a diagnosis of CHF or COPD.  If the diagnosis field has CHF or COPD in it, then the patient will have a care coordinator assigned.

For this exercise, write another function that takes the same inputs as LACE() plus another diagnosis parameter, and return True or False depending on if the paient needs a care coordinator.

**NOTE** Pay attention to the fact that the order of parameters in this function definition are not the same as the order they were in the LACE score.  

```python
def diagnosis_cd_score(diagnosis_cd):
        if diagnosis_cd == "CHF":
            return "CHF"
        elif diagnosis_cd == "COPD": 
            return "COPD"
        else:
            return "None"
        
def ed_visits_score(ed_visits):
        if ed_visits == 0:
            return 0
        elif ed_visits == 1:
            return 1
        elif ed_visits == 2:
            return 2
        elif ed_visits == 3:
            return 3
        else:
            return 4
        
def LOS_score(length_of_stay):
        if length_of_stay < 1:
            return 0
        elif length_of_stay == 1:
            return 1
        elif length_of_stay == 2:
            return 2
        elif length_of_stay == 3:
            return 3
        elif length_of_stay >= 4:
            return 4
        elif length_of_stay >= 7:
            return 5
        else:
            return 7
        
def acute_flag_score(acute_flag):
        if acute_flag:
            return 3
        else:
            return 0 
    
def charlson_score(charlson):
        if charlson == 0:
            return 0
        elif charlson == 1:
            return 1
        elif charlson == 2:
            return 2
        elif charlson == 3:
            return 3
        else:
            return 5
        
```

<!-- grade_id: weight-answer, points: 0, grade: false, solution: true -->
```python
def assign_care_coordinator(diagnosis_cd, ed_visits, length_of_stay, acute_flag, charlson):    
    if (ed_visits + length_of_stay + acute_flag + charlson) > 10:
        return True
    elif diagnosis_cd == "None":
        return False
    elif diagnosis_cd == "CHF":
        return True
    elif diagnosis_cd == "COPD":
        return True
    else:
        return none
    
    
    ### BEGIN SOLUTION
        assignment = None
    ### END SOLUTION
        
        return assignment
    
```

```python
assign_care_coordinator('None', 0, 4, False, 1)
```

Output:
```
False
```

```python
assign_care_coordinator('CHF', 0, 4, False, 1)
```

Output:
```
True
```

```python
assign_care_coordinator('COPD', 0, 4, False, 1)
```

Output:
```
True
```

```python
assign_care_coordinator('None', 7, 4, True, 4)
```

Output:
```
True
```

```python
assign_care_coordinator('CHF', 7, 4, True, 4)
```

Output:
```
True
```

<!-- grade_id: weight-tests, points: 5, grade: true, solution: false -->
```python
assert assign_care_coordinator('None', 0, 4, False, 1) == False
assert assign_care_coordinator('CHF', 0, 4, False, 1) == True
assert assign_care_coordinator('COPD', 0, 4, False, 1) == True
assert assign_care_coordinator('None', 7, 4, True, 4) == True
assert assign_care_coordinator('CHF', 7, 4, True, 4) == True
```