# Working with multiple data sets

There are two data files that we'll be working with for this week's assignment.  They are described below.  Load those data files in with Pandas and then work to answering each of the questions below.

## npidata.csv

This file is basic information about every healthcare provider in the US.  It has one row for each NPI (National Provider Identifier).  It contains information such as the provider's name and address.


## cmsYYYY.csv

These are files about what kinds of procedures and patients providers in the US are serving under CMS programs, Medicare and Medicaid.  Each contains various statistics about providers over the course of a year.  There are three of these, for the years 2014, 2015, and 2016.  These files, however, don't contain information about the provider such as where the provider is located.

## Our Goals

For this assignment, we're going to want to compute some statistics based on the data in the **cms** files, but aggregate that data based on information in the **npidata** file.  As in last week's assignment, you'll need to store your answers in a variable called `answer` at the end of each step.

```python
import pandas as pd
npi_data = pd.read_csv('/data/npidata.csv')
cms_2014 = pd.read_csv('/data/cms2014.csv')
cms_2015 = pd.read_csv('/data/cms2015.csv')
cms_2016 = pd.read_csv('/data/cms2016.csv')
```

## Part 1

In this first step, we'll need to merge together all of the **cms** files into a single dataframe.  Be careful that these files might not be identical, so you'll have to look a little bit to figure out how to merge them.

As you are merging them, make sure that you retain information about which file (i.e. which year) the data came from.  Call that new columns `year`.

In your `answer` variable, provide a complete data frame that contains all of the rows and columns from the **cms** files, plus an additional column to store the year/file that particular row came from.

The assertion tests will give you a good idea as to if you're merging the files correctly.

<!-- grade_id: 01-solution, points: 0, grade: false, solution: true -->
```python
import os, glob
import pandas as pd 
cms_2014 = pd.read_csv('/data/cms2014.csv')
cms_2015 = pd.read_csv('/data/cms2015.csv')
cms_2016 = pd.read_csv('/data/cms2016.csv')

answer = None

### SOLUTION

path = "/home/hds5210-2020/mydata"
all_files = glob.glob(os.path.join(path, "*.csv"))
all_df = []
for f in all_files:
    df = pd.read_csv(f, sep=",")
    f['file'] = f.split('/')[-1]
    all_df.append(df)
    
answer = pd.concat(all_df, ignore_index=True, sort=True)
answer["year"] = ['2014', '2015', '2016']

```

Output:
```
/opt/tljh/user/lib/python3.6/site-packages/pandas/core/reshape/concat.py in concat(objs, axis, join, join_axes, ignore_index, keys, levels, names, verify_integrity, sort, copy)
    253         verify_integrity=verify_integrity,
    254         copy=copy,
--> 255         sort=sort,
    256     )
    257 

/opt/tljh/user/lib/python3.6/site-packages/pandas/core/reshape/concat.py in __init__(self, objs, axis, join, join_axes, keys, levels, names, ignore_index, verify_integrity, copy, sort)
    302 
    303         if len(objs) == 0:
--> 304             raise ValueError("No objects to concatenate")
    305 
    306         if keys is None:

ValueError: No objects to concatenate
```

<!-- grade_id: 01-tests, points: 5, grade: true, solution: false -->
```python
assert(answer.shape == (193862, 60))
assert(list(answer['year'].unique()) == [2014,2015,2016])
assert(set(answer.columns.str.lower()) == set(['year', 'nbr', 'npi', 'provider_type',
       'medicare_participation_indicator', 'number_of_hcpcs', 'total_services',
       'total_unique_benes', 'total_submitted_chrg_amt',
       'total_medicare_allowed_amt', 'total_medicare_payment_amt',
       'total_medicare_stnd_amt', 'drug_suppress_indicator',
       'number_of_drug_hcpcs', 'total_drug_services',
       'total_drug_unique_benes', 'total_drug_submitted_chrg_amt',
       'total_drug_medicare_allowed_amt', 'total_drug_medicare_payment_amt',
       'total_drug_medicare_stnd_amt', 'med_suppress_indicator',
       'number_of_med_hcpcs', 'total_med_services', 'total_med_unique_benes',
       'total_med_submitted_chrg_amt', 'total_med_medicare_allowed_amt',
       'total_med_medicare_payment_amt', 'total_med_medicare_stnd_amt',
       'beneficiary_average_age', 'beneficiary_age_less_65_count',
       'beneficiary_age_65_74_count', 'beneficiary_age_75_84_count',
       'beneficiary_age_greater_84_count', 'beneficiary_female_count',
       'beneficiary_male_count', 'beneficiary_race_white_count',
       'beneficiary_race_black_count', 'beneficiary_race_api_count',
       'beneficiary_race_hispanic_count', 'beneficiary_race_natind_count',
       'beneficiary_race_other_count', 'beneficiary_nondual_count',
       'beneficiary_dual_count', 'beneficiary_cc_afib_percent',
       'beneficiary_cc_alzrdsd_percent', 'beneficiary_cc_asthma_percent',
       'beneficiary_cc_cancer_percent', 'beneficiary_cc_chf_percent',
       'beneficiary_cc_ckd_percent', 'beneficiary_cc_copd_percent',
       'beneficiary_cc_depr_percent', 'beneficiary_cc_diab_percent',
       'beneficiary_cc_hyperl_percent', 'beneficiary_cc_hypert_percent',
       'beneficiary_cc_ihd_percent', 'beneficiary_cc_ost_percent',
       'beneficiary_cc_raoa_percent', 'beneficiary_cc_schiot_percent',
       'beneficiary_cc_strk_percent', 'beneficiary_average_risk_score']))
```

Output:
```
AttributeError                            Traceback (most recent call last)
<ipython-input-6-bdb8e920e30d> in <module>
----> 1 assert(answer.shape == (193862, 60))
      2 assert(list(answer['year'].unique()) == [2014,2015,2016])
      3 assert(set(answer.columns.str.lower()) == set(['year', 'nbr', 'npi', 'provider_type',
      4        'medicare_participation_indicator', 'number_of_hcpcs', 'total_services',
      5        'total_unique_benes', 'total_submitted_chrg_amt',

AttributeError: 'NoneType' object has no attribute 'shape'
```

```python

```

## Part 2

In this next part, we're going to join the **cms** data with the provider information in the **/data/npidata.csv** file.  In this join, we don't want to lose any records from the **cms** files, even if no matching provider exists in the **npidata** file.  However, we don't care about any providers from the **npidata** file that don't have records in the **cms** files.  Those providers can be ignored.

Join the data files together to create one unified data frame called `answer`.  This dataframe should have all the columns from both **cms** and **npidata** files, joined together using the `npi` column.  

Note that `npi` is unique in the **npidata** file.

<!-- grade_id: 02-solution, points: 0, grade: false, solution: true -->
```python
answer = None
import os, glob
import pandas as pd

### SOLUTION
path = "/data/npidata.csv"
file = glob.glob(os.path.join(path, "*.csv"))
df = []
for f in file:
    df = pd.read_csv(f, sep=",")
    f['file'] = f.split('/')[-1]
    df.append(df)

path = "/home/hds5210-2020/mydata"
all_files = glob.glob(os.path.join(path, "*.csv"))
all_df = []
for f in all_files:
    df = pd.read_csv(f, sep=",")
    f['file'] = f.split('/')[-1]
    all_df.append(df)
    
answer = pd.concat(all_df, df, ignore_index=True, sort=True)



```

Output:
```
/opt/tljh/user/lib/python3.6/site-packages/pandas/core/reshape/concat.py in concat(objs, axis, join, join_axes, ignore_index, keys, levels, names, verify_integrity, sort, copy)
    253         verify_integrity=verify_integrity,
    254         copy=copy,
--> 255         sort=sort,
    256     )
    257 

/opt/tljh/user/lib/python3.6/site-packages/pandas/core/reshape/concat.py in __init__(self, objs, axis, join, join_axes, keys, levels, names, ignore_index, verify_integrity, copy, sort)
    302 
    303         if len(objs) == 0:
--> 304             raise ValueError("No objects to concatenate")
    305 
    306         if keys is None:

ValueError: No objects to concatenate
```

<!-- grade_id: 02-tests, points: 5, grade: true, solution: false -->
```python
import numpy
assert(answer.shape == (193862, 103))
assert(list(answer['Provider Business Mailing Address State Name'].unique()) == ['IL','MO',numpy.nan,'WY'])
assert(list(answer.groupby('Provider Business Mailing Address State Name').NPI.count()) == [111520, 53366, 4805])
```

Output:
```
AttributeError                            Traceback (most recent call last)
<ipython-input-4-6539db35ef6b> in <module>
      1 import numpy
----> 2 assert(answer.shape == (193862, 103))
      3 assert(list(answer['Provider Business Mailing Address State Name'].unique()) == ['IL','MO',numpy.nan,'WY'])
      4 assert(list(answer.groupby('Provider Business Mailing Address State Name').NPI.count()) == [111520, 53366, 4805])

AttributeError: 'NoneType' object has no attribute 'shape'
```

## Part 3

If you did everything right above, you'll notice that grouping by a column with NaN in it will cause some rows to disappear from the aggregation test.  So, let's create a new column called `'State'` that has the same value as whatever is in the `'Provider Business Mailing Address State Name'` column or a value of `'XX'` if there is no state information.

Set `answer` to be your final data frame with the new `'State'` column added.

<!-- grade_id: 03-solution, points: 0, grade: false, solution: true -->
```python
answer = None

### SOLUTION

answer["State"] = "Provider Business Mailing"

import numpy as np
df.replace(np.nan, "XX")


```

Output:
```
TypeError                                 Traceback (most recent call last)
<ipython-input-5-8728b19d7001> in <module>
      3 ### SOLUTION
      4 
----> 5 answer["State"] = "Provider Business Mailing"
      6 
      7 import numpy as np

TypeError: 'NoneType' object does not support item assignment
```

<!-- grade_id: 03-tests, points: 5, grade: true, solution: false -->
```python
assert(list(answer.groupby('State').npi.count()) == [111520, 53366, 4805, 24171])
assert(answer.shape == (193862, 104))
```

Output:
```
AttributeError                            Traceback (most recent call last)
<ipython-input-8-ead5e4353622> in <module>
----> 1 assert(list(answer.groupby('State').npi.count()) == [111520, 53366, 4805, 24171])
      2 assert(answer.shape == (193862, 104))

AttributeError: 'NoneType' object has no attribute 'groupby'
```

## Part 4

Next, let's summarize the data by year and by State.  Create a pivot table that contains one row for each state and one column for each year.  Within the pivot table, put a sum of total services as the values.

Assign `answer` to be that resulting pivot table.  In the tests, I'm going to plot a bar chart of your pivot table.

<!-- grade_id: 04-solution, points: 0, grade: false, solution: true -->
```python
%matplotlib inline
answer = None

### SOLUTION
table = pd.pivot_table(answer)

```

Output:
```
<ipython-input-6-681554015a0c> in <module>
      3 
      4 ### SOLUTION
----> 5 table = pd.pivot_table(answer)

/opt/tljh/user/lib/python3.6/site-packages/pandas/core/reshape/pivot.py in pivot_table(data, values, index, columns, aggfunc, fill_value, margins, dropna, margins_name, observed)
     85 
     86     else:
---> 87         values = data.columns
     88         for key in keys:
     89             try:

AttributeError: 'NoneType' object has no attribute 'columns'
```

<!-- grade_id: 04-tests, points: 5, grade: true, solution: false -->
```python
assert(answer.shape == (4,3))
assert(answer.sum().sum() == 519185664.6999999)
assert(answer[2016].sum() == 176596933.80000004)
assert(answer.loc['WY'].sum() == 10892707.4)
```

Output:
```
AttributeError                            Traceback (most recent call last)
<ipython-input-9-db10226393b0> in <module>
----> 1 assert(answer.shape == (4,3))
      2 assert(answer.sum().sum() == 519185664.6999999)
      3 assert(answer[2016].sum() == 176596933.80000004)
      4 assert(answer.loc['WY'].sum() == 10892707.4)

AttributeError: 'NoneType' object has no attribute 'shape'
```

```python
%matplotlib inline
answer.plot.bar()
```

Output:
```
AttributeError                            Traceback (most recent call last)
<ipython-input-10-5fc75f4fb76d> in <module>
      1 get_ipython().run_line_magic('matplotlib', 'inline')
----> 2 answer.plot.bar()

AttributeError: 'NoneType' object has no attribute 'plot'
```

```python

```