# Working with a database

This week's assignment has a few basic steps.  First, we're going to pull some data down off the internet and store it into our MySQL database.  Make sure that you use your username as part of the table name as show in the examples so that you don't create a problem for other students.

Then, we'll merge that with some data already in the database and calculate a few results.  When it comes to calculating the results, you can do so either with SQL or with Pandas operations.


## PART 1: Setup your database connection and table name

In the code below, change the value of the variable `MYTABLE` to use your own username rather that `'pboal'`

You can then use `MYTABLE` in the rest of your code to reference that table name.

<!-- grade_id: q01-solution, points: 0, grade: false, solution: true -->
```python
#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!
# First, put in your user name below
#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!#!

MYTABLE = 'sdraper16' + '_data'

import os
import pymysql
from sqlalchemy import create_engine
import pandas as pd


host = 'slucor2020.cgdcoitnku0k.us-east-1.rds.amazonaws.com'
port = '3306'
user = 'slucor2020'
password = 'SLUcor2020'
database = 'hds5210'

conn = create_engine('mysql+pymysql://' + 
                     user + 'sdraper16' + 
                     password + 'SLUcor2020' + 
                     host + 'slucor2020.cgdcoitnku0k.us-east-1.rds.amazonaws.com' + 
                     database, echo=False)

print(conn)

if(conn):
    print("Connection is succesfull")
else:
    print("Connection failed")

# Then, create a connection to the same database, using the 
# same user name and password in the `mysql-examples.ipynb` 
# file.  Use conn as the database connection variable.

conn = None



```

Output:
```
Engine(mysql+pymysql://slucor2020sdraper16SLUcor2020SLUcor2020slucor2020.cgdcoitnku0k.us-east-1.rds.amazonaws.comslucor2020.cgdcoitnku0k.us-east-1.rds.amazonaws.comhds5210)
Connection is succesfull
```

<!-- grade_id: q01-tests, points: 5, grade: true, solution: false -->
```python
assert(MYTABLE != 'sdraper16_data')
assert(conn.connect())
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-4-d06132d25b10> in <module>
----> 1 assert(MYTABLE != 'sdraper16_data')
      2 assert(conn.connect())

AssertionError:
```

## PART 2: Bring in outside data

Grab the data from this URL and put it into a database table named with your `username_data`.

http://dhcs-chhsagency.opendata.arcgis.com/datasets/8e4f3a0c75b9424d888d11c1f949cc32_0.csv

By the end of your cell, the table should be created.  The tests are going to verify that the table exists and looks right.

<!-- grade_id: q02-solution, points: 0, grade: false, solution: true -->
```python

### SOLUTION GOES HERE

import MySQLdb as mdb
connection = mdb.connect("slucor2020.cgdcoitnku0k.us-east-1.rds.amazonaws.com", "3306", "sdraper16", "hds5210")
cursor = connection.cursor()
url = mdb.escape_string('http://dhcs-chhsagency.opendata.arcgis.com/datasets/8e4f3a0c75b9424d888d11c1f949cc32_0.csv')
cursor.execute("""INSERT INTO "sdraper16_data"(slucor2020.cgdcoitnku0k.us-east-1.rds.amazonaws.com) VALUES(%s)""", (url,))

for sdraper16_data in result:
    print(sdraper16_data)
```

<!-- grade_id: q02-tests, points: 5, grade: true, solution: false -->
```python
dxyz = pd.read_sql_query('SELECT * FROM ' + MYTABLE, conn)
assert(dxyz.shape == (328,22))
assert(list(dxyz.columns) == ['index', 'X', 'Y', 'Provider_Name', 'NPI', 'CCN',
       'Medicaid_EP_Hospital_Type', 'Street_Address', 'City', 'County',
       'State', 'Zip_Code', 'Payment_Year_Number', 'Program_Type',
       'Total_payments', 'Last_Program_Year', 'Last_Payment_Year',
       'Last_Payment_Criteria', 'Most_Recent_Disbursement_Amount', 'LONGITUDE',
       'LATITUDE', 'FID'])
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-1-360ae4defa77> in <module>
----> 1 dxyz = pd.read_sql_query('SELECT * FROM ' + MYTABLE, conn)
      2 assert(dxyz.shape == (328,22))
      3 assert(list(dxyz.columns) == ['index', 'X', 'Y', 'Provider_Name', 'NPI', 'CCN',
      4        'Medicaid_EP_Hospital_Type', 'Street_Address', 'City', 'County',
      5        'State', 'Zip_Code', 'Payment_Year_Number', 'Program_Type',

NameError: name 'pd' is not defined
```

## PART 3: Combine with other data in the database

In the database, there is an existing table called `population`.  We want to merge the DHCS datafile loaded above with the population data available in this other database table  The tables can be merged on `MYTABLE`'s `Zip_Code` field and `population`'s `zipcode` field.

Note that not all `Zip_Codes` from your downloaded file have to be in the `population` table.  If they aren't, then I want you to eliminate the non-matching records.  That is, only keep the records that have a matching ZIP code in both sets of data.

Answer the question:
Which providers are located in the zipcode with the largest population?

Put your answer in the form `answer = ['a', 'list', 'of', 'NPI', 'like', '1593042103]`

<!-- grade_id: q03-solution, points: 0, grade: false, solution: true -->
```python
### SOLUTION GOES HERE
answer = []

import mysql.connector
import pandas as pd

url = 'http://dhcs-chhsagency.opendata.arcgis.com/datasets/8e4f3a0c75b9424d888d11c1f949cc32_0.csv'
answer = pd.read_sql_query('SELECT * FROM ' + MYTABLE, conn)

answer = []
data = [
    ['Index', 16],
    ['Provider_Name',221],
    ['zipcode',245],
    ['LONGITUDE', 1194016923]
]
df = pd.DataFrame(data, answer=answer)

```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-2-d8a5b5555665> in <module>
      6 
      7 url = 'http://dhcs-chhsagency.opendata.arcgis.com/datasets/8e4f3a0c75b9424d888d11c1f949cc32_0.csv'
----> 8 answer = pd.read_sql_query('SELECT * FROM ' + MYTABLE, conn)
      9 
     10 answer = []

NameError: name 'MYTABLE' is not defined
```

<!-- grade_id: q03-tests, points: 5, grade: true, solution: false -->
```python
assert(type(answer) == list)
assert(answer == ['1194016923'])
```

## PART 4: Total by hospital type

This is a multistep process:
* From our downloaded data file, compute the `Total payments` per ZIP code and Medicaid EP Hospital Type.
* Then merge that with the `population` data to compute a `Total payments` per person.
* Then average that across all of the `Medicaid EP Hospital Types` to get an average per persona payment for these type of hospital.

Your answer should be in structure of a data frame with at least two columns:
* Medicaid_EP_Hospital_Type
* Avg_Pay_per_Capita

<!-- grade_id: q04-solution, points: 0, grade: false, solution: true -->
```python
### SOLUTION GOES HERE
url = 'http://dhcs-chhsagency.opendata.arcgis.com/datasets/8e4f3a0c75b9424d888d11c1f949cc32_0.csv'
answer=pd.read_sql_query("""
  SELECT 
    h.*, v.year, 
    CASE WHEN Total payments ELSE NULL END as per person,
    CASE WHEN hospital_services THEN Medicaid_EP_Hospital_Type ELSE NULL END as Acute Care Hospitals
  FROM 
    population,
    hospital_services
  WHERE
    Medicaid_EP_Hospital_Type per person,
    Avg_Pay_per_Capita
""", conn)

answer = None
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-3-eb892cc444e3> in <module>
     12     Medicaid_EP_Hospital_Type per person,
     13     Avg_Pay_per_Capita
---> 14 """, conn)
     15 
     16 answer = None

NameError: name 'conn' is not defined
```

<!-- grade_id: q04-tests, points: 5, grade: true, solution: false -->
```python
assert(type(answer) == pd.core.frame.DataFrame)
assert(round(answer.query("Medicaid_EP_Hospital_Type == 'Acute Care Hospitals'")['Avg_Pay_per_Capita'].sum(),3) == 17974.363)
assert(round(answer.query("Medicaid_EP_Hospital_Type == 'Children\\'s Hospitals'")['Avg_Pay_per_Capita'].sum(),3) == 111.559)

```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-4-c2c725d8c450> in <module>
----> 1 assert(type(answer) == pd.core.frame.DataFrame)
      2 assert(round(answer.query("Medicaid_EP_Hospital_Type == 'Acute Care Hospitals'")['Avg_Pay_per_Capita'].sum(),3) == 17974.363)
      3 assert(round(answer.query("Medicaid_EP_Hospital_Type == 'Children\\'s Hospitals'")['Avg_Pay_per_Capita'].sum(),3) == 111.559)

AssertionError:
```