# Week 6 Exercises

_McKinney 6.1_

There are multiple ways to solve the problems below.  You can use any one of several approaches.  For example, you can read CSV files using Pandas or the csv module.  Your score won't depend on which modules you choose to use unless explicitly noted below, but your programming style will still matter.

### 26.1 - When a Python program opens a file to read from it, no other programs can read from the same file at the same time.

> True or False


```python
answer261 = False
```

### 26.2 - What is the advantage of opening files in Python using a `with` block?

> A) It ensures the file is closed as soon as you're done using it
>
> B) It runs faster than not using the `with` block
>
> C) It automatically loads the whole file into memory so that it can be accessed faster
>
> D) Actually, it is the only way to do it


```python
answer262 = "A"
```

### 26.3 - Which of the following is true?

> A) Many files only contain text, but many file formats do not use text
>
> B) All files contain only text, it's the just the text contents that explain the data to the computer
>
> C) All files contain only text in either JSON, XML, or CSV.
>
> D) Even if files contain only text, there still has to be a special code at the beginning to explain that this file will contain only text.


```python
answer263 = "A"
```

### 27.1 - Which of the following do you think is least likely to be used as a delimiter in structured files?

> A) `,`
>
> B) `|`
>
> C) `[tab]`
>
> D) `[space]`


```python
answer271 = "C"
```

### 27.2 - Which approaches can be used to read pipe-delimited (`|`) files?

*Report your answer as a list of options, e.g.: `["A","C"]`*

> A) Using the csv module
>
> B) Using the pandas module
>
> C) Using built-in file reading
>
> D) Using the psv module


```python
answer272 = ["A","B"]
```

### 27.3 - In a CSV file, how does the format change if you want to have a `,` in the actual contents of the file rather than as a delimiter?

> A) You can't do that.  You have to remove every occurence of `,`s
>
> B) Enclose the `,` using quotation marks
>
> C) Escape the `,` using a back-slash
>
> D) Nothing special, it will be obvious to the reader


```python
answer273 = "C"
```

### 28.1 - Select all of the data types that can be represented within a JSON format:

Provide your answer in the form of a list.  For example: `answer = ["A","B"]` 

> A) str
>
> B) list
>
> C) dict
>
> D) int


```python
answer281 = ["A"."D"]
```

### 28.2 - JSON stands for:

> A) Just a Simple Object Notation
>
> B) JavaScript Object Notation
>
> C) JavaScript Objects and Numbers
>
> D) Java Standard Object Notation


```python
answer282 = "B"
```

### 28.3 - Without just running the code, what will the following code print?

```
import json

data = """
[  {"id": { "last name": "Boal", "first name": "Paul", "middle initial": "E" }},
   {"id": { "last name": "Westhus", "first name": "Eric" }} ]
"""

names = json.loads(data)
list(names[1]['id'].keys())[1]
```

> A) 'Eric'
>
> B) 'first name'
>
> C) 'last name'
>
> D) 'Boal'


```python
answer283 = "B"
```

### 29.1 - In this class, XML Stands for:

> A) eXtensible Markup Language
>
> B) The 10th version of the Markup Language
>
> C) eXhange Markup Language
>
> D) X-Men Legends


```python
answer291 = "A"
```

### 29.2 - Which of the following is not a type of data in an XML document?

> A) tags
>
> B) attributes
>
> C) text
>
> D) list


```python
answer292 = "D"
```

### 29.3 - True or False, each element in an XML document can only have one child with a specific tag.

For example, is the following XML illegal because there are two child elements with the tag "kid":

```xml
<family>
  <parent role="dad">Paul Boal</parent>
  <kid role="daughter">Eleanor</kid>
  <kid role="daughter">Adelaide</kid>
</family>
```

Answer `False` if you think the above XML is valid.  

Answer `True` if you think the above XML is invalid.


```python
answer293 = False 
```

### 30.1 List of Allergies

In this directory, there is a file called `allergies.json` that contains a list of patient allergies.  It is taken from sample data provided by the EHR vendor, Epic, here: https://open.epic.com/Clinical/Allergy

Take some time to look at the structure of the file.  You can open it directly in Jupyter.

Within the file, you'll see that it is a dictionary with many items in it.  One of those items is called `entry` and that item is a list of things.  You can tell that because the item name is immediately followed by an opening square bracket, signifying the start of a list.  It's line 11 of the file: `  "entry": [`

Write a function named `allergy_count(json_file)` that takes as one parameter the name of the JSON file and returns an integer number of entries in that file.  Your function should open the file, read the json into a Python object, and return how many items there are in the list of `entry`s.

```python
import json

### BEGIN SOLUTION
def allergy_count(json_file): 
    allergies = json.load(open('allergies.json'))
    
entries_in_file = 0

for entry in allergies:
    entries_in_file = entry.get('entries')
    
print(entries_in_file)
### END SOLUTION
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-1-79726fe21610> in <module>
      7 entries_in_file = 0
      8 
----> 9 for entry in entries:
     10     entries_in_file = entry.get('Entries')
     11 

NameError: name 'entries' is not defined
```

<!-- grade_id: q1-tests, points: 5, grade: true, solution: false -->
```python
assert type(allergy_count('allergies.json')) == int
assert allergy_count('allergies.json') == 3

```

Output:
```
<ipython-input-6-0627ce555edf> in <module>
----> 1 assert type(allergy_count('allergies.json')) == int
      2 assert allergy_count('allergies.json') == 3

<ipython-input-5-e7047e57f5d7> in allergy_count(json_file)
     10 
     11     json_file = 'data/' + json_file
---> 12     entry = json.load(open('allergies.json'))
     13     return len(entry['entry'])
     14     print(entry)

FileNotFoundError: [Errno 2] No such file or directory: 'allergies.json'
```

### 30.2 Patient Allergies and Reaction

You'll see in the file that each of the items in the `entry` list have several other attributes including a patient name, substance text representation, and a reaction manifestation.  Create a function named `allergy_list(json_file)` that will create an output list that has patient name, allergy, and reaction for each `entry`.  The actual result you should get will be:

```python
[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis']]
```

You'll notice that the reaction and the manifestation of that action are lists.  You only need to capture the first reaction and the first manifestation of the action.  That is, if there is a list of things, just output the first one.

<!-- grade_id: q2-answer, points: 0, grade: false, solution: true -->
```python
import json

### BEGIN SOLUTION
def allergy_list(json_file):
    json_file = 'data/' + json_file
    entries = json.load(open('allergies.json'))
    
    all_entries = [] 
    for entry in entries ['entry']:
        patient = entry['resource']['patient']['display']
        allergy = entry['resource']['allergy']['text']
        reaction = entry['resource']['reaction'][0]['text']
        
        all_entries.append([patient, allergy, reaction])
    
    return all_entries
   
    
### END SOLUTION
```

<!-- grade_id: q2-tests, points: 10, grade: true, solution: false -->
```python
output=[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis']]

assert allergy_list('allergies.json') == output

```

Output:
```
<ipython-input-3-3c58096b8196> in <module>
      3  ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis']]
      4 
----> 5 assert allergy_list('allergies.json') == output

<ipython-input-2-02aea1c8fa71> in allergy_list(json_file)
      4 def allergy_list(json_file):
      5     json_file = 'data/' + json_file
----> 6     entries = json.load(open(json_file))
      7 
      8     all_entries = [] #create empty list

FileNotFoundError: [Errno 2] No such file or directory: 'data/allergies.json'
```

### 30.3 Allergy Reaction

Write a function called `allergy_reaction(json_file,substance)` that takes two parameter and returns the reaction that will happen if the patient takes the specified substance.  Solve this, in part, by calling your `allergy_list` function inside your new `allergy_reaction` function.

If the substance is not found in the allergy list, the function should return None.

<!-- grade_id: q3-answer, points: 0, grade: false, solution: true -->
```python
import json

### BEGIN SOLUTION

def allergy_reaction(json_file,substance): 
    all_entries = allergy_list(json_file)
    entries = json.load(open('allergies.json'))
    
    for entry in all_entries:
        if entry== 'PENICILLIN G':
            print('Hives')
        elif entry == 'SHELLFISH-DERIVED PRODUCTS':
            print('Itching')
        elif entry == 'STRAWBERRY':
            print('Ananphylaxis')
        else:
            print(None) 
        
### END SOLUTION
```

<!-- grade_id: q3-tests, points: 5, grade: true, solution: false -->
```python
assert allergy_reaction('allergies.json','PENICILLIN G') == 'Hives'
assert allergy_reaction('allergies.json','SHELLFISH-DERIVED PRODUCTS') == 'Itching'
assert allergy_reaction('allergies.json','STRAWBERRY') == 'Anaphylaxis'
assert allergy_reaction('allergies.json','PENICILLIN') == None
```

Output:
```
<ipython-input-9-b85182145fad> in allergy_reaction(json_file, substance)
      4 
      5 def allergy_reaction(json_file,substance):
----> 6     all_entries = allergy_list(json_file)
      7     entries = json.load(open('allergies.json'))
      8 

<ipython-input-6-2ac96d5eccb8> in allergy_list(json_file)
      4 def allergy_list(json_file):
      5     json_file = 'data/' + json_file
----> 6     entries = json.load(open('allergies.json'))
      7 
      8     all_entries = []

FileNotFoundError: [Errno 2] No such file or directory: 'allergies.json'
```

### STRETCH - For those looking for an additional challenge

The Coronavirus is creating quite the stir right now.  There are some sources suggesting that trends show it is going to be significantly more serious than SARS was back in the 2002 timeframe.  Here's one visualization trying to demonstrate that: https://www.reddit.com/r/China_Flu/comments/ev2b4v/i_updated_some_charts_comparing_this_outbreak/

Someone on Kaggle has generously already compiled a dataset based on information from Johns Hopkins about the Coronavirus outbreak.  https://www.kaggle.com/brendaso/2019-coronavirus-dataset-01212020-01262020  Create a Kaggle account, if you don't already have one.  Download this data set and then upload it to your Jupyter Home folder.  (The "up arrow" button is for uploading a file.)

Use Python's built-in `csv` module to read the data from this file and generate the following information: **what are the total confirmed cases in all of Mainland China as of the latest information in the data set?**  Some important things to note:
* Each entry for a given city has the **cumulative** number of cases.  So that column is not additive (it cannot be summed).  You'll have to find a way to filter your data for the last day for each city, then total those up.
* If you choose to parse the date column, you will want to lookup how to do that using Python's `datetime` module.  Especially the `strptime` function.  https://docs.python.org/3/library/datetime.html#strftime-strptime-behavior  Hint: you can parse a date string in the format 2/17/2020 using the code below.  This link will tell you what things like `%m` and `%Y` mean.

```
from datetime import datetime
d = datetime.strptime('2/17/2020', '%m/%d/%Y')
```

If you want to take this another step, **create a list of tuples that contain (observate date, total confirmed) totalled over all locations represented in the data**

```python

```

```python

```